## Build a model using Keras

```python
import pandas as pd
import numpy as np
import matplotlib.pyplot as plt
import requests

```

```python
### Download data from google drive. You need not mess with this code.

import requests

def download_file_from_google_drive(id, destination):
    URL = "https://docs.google.com/uc?export=download"

    session = requests.Session()

    response = session.get(URL, params = { 'id' : id }, stream = True)
    token = get_confirm_token(response)

    if token:
        params = { 'id' : id, 'confirm' : token }
        response = session.get(URL, params = params, stream = True)

    save_response_content(response, destination)    

def get_confirm_token(response):
    for key, value in response.cookies.items():
        if key.startswith('download_warning'):
            return value

    return None

def save_response_content(response, destination):
    CHUNK_SIZE = 32768

    with open(destination, "wb") as f:
        for chunk in response.iter_content(CHUNK_SIZE):
            if chunk: # filter out keep-alive new chunks
                f.write(chunk)
                
if __name__ == "__main__":
    file_id = '1DHF4b0sBB_KLQ4oxNEMp0sGrViu0gpeG'
    destination = 'data.csv'
    download_file_from_google_drive(file_id, destination)
```

```python
import pandas as pd
import numpy as np
import matplotlib.pyplot as plt
import requests


# Importing and cleaning data using pandas library
data = pd.read_csv('data.csv')

## Last column is output features. Rest are inputs
X = data.iloc[:, 1:-1].values

# 2nd column is output labels
y = data.iloc[:, -1].values
```

Convert labels to intermediate representation where each label is replaced by a number from 0 to Ny-1
(Ny is number of classes)

<!-- grade_id: cell-a674018e3af0b38c, points: 0, grade: false, solution: true -->
```python
def conv_labels(y_raw):
    """
    Inputs:
        y_raw: numpy array of labels
    Outputs:
        y: numpy array of ints, each label is replaced by an int from 0 to Ny-1
        Ny: number of classes
    """
    # YOUR CODE HERE
```

<!-- grade_id: cell-c59e63d3a5ab1587, points: 3, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-18c5d990cc8c0b88, points: 0, grade: false, solution: true -->
```python
### One-hot encode Y_v
def oneHot(y, Ny):
    '''
    Input:
        y: an int in {0, 1}
        Ny: Number of classes, e.g., 2 here.
    Output:
        Y: a vector of Ny (=2) tuples
    '''
    # YOUR CODE HERE
```

<!-- grade_id: cell-63afef22a1950397, points: 2, grade: true, solution: false -->
```python
"""Test for one-hot"""
assert np.all(oneHot(0,3)==np.array([1,0,0]))

```


#### Create Y_o which is one-hot encoding of Y using above functions


<!-- grade_id: cell-35e2ebfbcce2b63f, points: 0, grade: false, solution: true -->
```python
def create_Y_o(y):
    """
    Inputs:
        y: numpy array of class labels
    Outputs:
        Y_o: numpy array of shape(samples, Ny) with one-hot encodings of y
        Ny: number of unique classes
    """
    # YOUR CODE HERE
    return Y_o, Ny

Y_o, Ny = create_Y_o(y)

```

<!-- grade_id: cell-5c2e4a519a112844, points: 2, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-87b2cacee71e7ed4, points: 0, grade: false, solution: true -->
```python
### Split data into train and test. Keep 10% of samples for testing
## Divide the data into these variables - X_train, X_test, y_train, y_test
# YOUR CODE HERE
```

<!-- grade_id: cell-cfa28681f99002d1, points: 2, grade: true, solution: false -->
```python
"""test for splitting"""
assert(X_train.shape[0] == 13608)
```

<!-- grade_id: cell-2bc42ab83e597a8e, points: 0, grade: false, solution: true -->
```python
## Normalize the Data
def findMeanStddev(X):
    '''
    Input: 
        X: a matrix of size (no. of samples, dimension of each sample)
    Output:
        mean: mean of samples in X (same size as X)
        stddev: element-wise std dev of sample in X (same size as X)
    '''
    # YOUR CODE HERE
```

<!-- grade_id: cell-085a9443367fdb41, points: 3, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-70997efec08284dd, points: 0, grade: false, solution: true -->
```python
def normalizeX(X, mean, stddev):
    '''
    Input:
        X: a matrix of size (no. of samples, dimension of each sample)
        mean: mean of samples in X (same size as X)
        stddev: element-wise std dev of sample in X (same size as X) 
    Output:
        Xn: X modified to have 0 mean and 1 std dev
    '''
    # YOUR CODE HERE

mean_train, stddev_train = findMeanStddev(X_train)
X_train = normalizeX(X_train, mean_train, stddev_train)
X_test = normalizeX(X_test, mean_train, stddev_train)
```

<!-- grade_id: cell-46adfce309e1b614, points: 2, grade: true, solution: false -->
```python

```

#### Create model. 
- Choose the number of hidden layers, neurons, activations, loss function, learning rate and optimizers on your own.
- Report accuracy metric
- Use no more than 100 epochs
- Use validation_split = 0.1

<!-- grade_id: cell-71762f0bb2dd5ac6, points: 0, grade: false, solution: true -->
```python
import keras
def create_model():
    """
    Inputs:
        None
    Outputs:
        model: keras model afteer compiling
    """
    # YOUR CODE HERE
    return model
```

```python
model = create_model()
history = model.fit(X_train, y_train, epochs=100, batch_size = 100, validation_split = 0.1)
```

```python
from matplotlib import pyplot as plt
plt.plot(history.history['val_acc'])
```

#### Evalutaion


<!-- grade_id: cell-2b2c6db73900f0cd, points: 8, grade: true, solution: false -->
```python
"""Test for model"""

```

<!-- grade_id: cell-1a1b44abeb716585, points: 1, grade: true, solution: false -->
```python
"""Test for model"""

```

<!-- grade_id: cell-8f0f28af43148408, points: 2, grade: true, solution: false -->
```python
"""Test for model"""

```

#### Confusion Matrix

<!-- grade_id: cell-6ab6b4c2b62d34e5, points: 0, grade: false, solution: true -->
```python
def create_cm(Y_test, Y_pred):
    """
    Inputs:
        Y_test: numpy array with true labels
        Y_pred: numpy array with predicted labels
    Outputs:
        CM: numpy array (ndim=2) containing confusion matrix
    """
    # YOUR CODE HERE
    return CM
```

<!-- grade_id: cell-689817aaae734a4e, points: 2, grade: true, solution: false -->
```python
"""Test for confusion matrix"""
```

<!-- grade_id: cell-eae6c93e269a7b33, points: 0, grade: false, solution: true -->
```python
def precision(cm, cls):
    """
    Inputs:
        cm: numpy array of shape (Ny, Ny) containing confusion matrix
        cls: class for which precision has to be computed
    Outputs:
        prec: precision for class cls 
    """
    # YOUR CODE HERE
```

<!-- grade_id: cell-abb3832f763effa7, points: 3, grade: true, solution: false -->
```python

```

```python

```